Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Exercise
Write a function that receives a list and returns it in reverse order every 2 elements, starting from the last-1 element

<!-- grade_id: cell-edb80791f1172d9c, points: 0, grade: false, solution: true -->
```python
def getreversed(l):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f39f060e7f0f3cd5, points: 0.2, grade: true, solution: false -->
```python
assert [2] == getreversed([1, 2, 3])
assert ['hola', 3, 1] == getreversed([1, 2, 3, 4 , "hola", "mundo"])
assert [3] == getreversed(getreversed([1, 2, 3, 4 , "hola", "mundo"]))
```

# Exercise
Write a function that receives a positive integer and returns a list with all its prime factors


<!-- grade_id: cell-985fbac2f9ea8a4b, points: 0, grade: false, solution: true -->
```python
import numpy as np

def primefactors(n):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def isprime(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a83025d2402f6fd4, points: 0.4, grade: true, solution: false -->
```python
assert [2] == primefactors(2)
assert [2, 5] == primefactors(10)
assert [2, 3] == primefactors(36)
assert [37] == primefactors(37)
assert [2, 19] == primefactors(38)
assert [2, 3, 5] == primefactors(90)
assert [10039] == primefactors(10039)
assert [2, 5, 251] == primefactors(10040)
assert [3, 3347] == primefactors(10041)
assert [1213427] == primefactors(1213427)
assert [2, 3, 101119] == primefactors(1213428)
```

# Exercise
Representing a polynomial by a dictionary : Consider the polynomial $p(x) = −1+x^2+3x^7$ . It can be represent by means of a dictionary as p = {0:-1, 2:1, 7:3} . Write a function which gets a dictionary representing a polynomial, an $x$ value, and returns the polynomial evaluated on that $x$ value. 

<!-- grade_id: cell-c8e849d0bd27febe, points: 0, grade: false, solution: true -->
```python
import numpy as np
def polynomial(coeffs, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fc4b987152ce277a, points: 0.4, grade: true, solution: false -->
```python
assert np.isclose(-1.0, polynomial({0:-1}, 0.2))
assert np.isclose(-1.0, polynomial({0:-1}, 0.57))
assert np.isclose(-0.96, polynomial({0:-1, 2:1}, 0.2))
assert np.isclose(-0.6751, polynomial({0:-1, 2:1}, 0.57))
assert np.isclose(-0.9599616, polynomial({0:-1, 2:1, 7:3}, 0.2))
assert np.isclose(-0.6164530752042101, polynomial({0:-1, 2:1, 7:3}, 0.57))
assert np.isclose(11.54503606479579, polynomial({0:10.98, 2:1, 3:0.98, 7:3}, 0.57))

```

# Exercise
Create a function that receives a numpy array and returns the maximum across a given axis. Hint: Check the `axis` parameter in the `.max` function.

<!-- grade_id: cell-5a0c7690be1be539, points: 0, grade: false, solution: true -->
```python
import numpy as np
def maximum(data, axis):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a4b3899af90c32f9, points: 0.2, grade: true, solution: false -->
```python
import numpy.random as random # do not change
random.seed(10) # do not change

array = np.array([1, 23.4, 3.0])
assert np.isclose(23.4, maximum(array, 0))

array = np.random.random((2, 3))
assert np.allclose(np.array([0.77132064, 0.49850701, 0.63364823]), maximum(array, 0))
assert np.allclose(np.array([0.77132064, 0.74880388]).reshape((1, 2)), maximum(array, 1))

array = np.random.random((2, 3))
assert np.allclose(np.array([0.19806286, 0.76053071, 0.95339335]), maximum(array, 0))
assert np.allclose(np.array([0.76053071, 0.95339335]).reshape((1,2)), maximum(array, 1))
```

# Exercise
Computing an integral using numpy: Write a function that receives the limits for an interval, $a$ and $b$, the number $n$ of intervals to use (which means $n+1$ points), and a function $f$, and returns the integral of that function over the interval using the simple trapezoidal rule

$$I(a, b, n, f) \simeq \Delta x \left( \frac{1}{2}f(x_0) + \frac{1}{2}f(x_n) + \sum\limits_{i=1}^{n-1}f(x_i) \right), $$

where $x_k = x_0 + k\Delta x$, $x_0 = a$, $x_n = b$, and $\Delta x = (b-a)/n$. Use `np.linspace`, `np.sum`, etc, and try to avoid using loops. 


<!-- grade_id: cell-194983d1b5284cb6, points: 0, grade: false, solution: true -->
```python
def trapezoid(a, b, n, f):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-40001aa792fe1bb7, points: 0.3, grade: true, solution: false -->
```python
assert np.isclose(0.7272870719562238, trapezoid(0.1, 1.3, 20, np.sin))
assert np.isclose(0.7275053148282778, trapezoid(0.1, 1.3, 2000, np.sin))
assert np.isclose(0.7275053366512558, trapezoid(0.1, 1.3, 200000, np.sin))
assert np.isclose(0.8634656357913549, trapezoid(0.1, 1.3, 20, np.cos))
assert np.isclose(0.8637247428586216, trapezoid(0.1, 1.3, 2000, np.cos))
assert np.isclose(0.8637247687677736, trapezoid(0.1, 1.3, 200000, np.cos))
```

```python

```