Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Write a function which receives a positive integer number and prints the sum of all even numbers smaller than the parameter read

<!-- grade_id: cell-48df8f2b83dee5a8, points: 0, grade: false, solution: true -->
```python
#import numpy
def sum_smaller_even(n):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4fe77a6d354d52e1, points: 0.5, grade: true, solution: false -->
```python
assert 0 == sum_smaller_even(-10)
assert 0 == sum_smaller_even(0)
assert 0 == sum_smaller_even(2)
assert 2 == sum_smaller_even(3)
assert 2 == sum_smaller_even(4)
assert 6 == sum_smaller_even(5)
assert 6 == sum_smaller_even(6)
assert 30 == sum_smaller_even(12)
assert 42 == sum_smaller_even(13)
assert 42 == sum_smaller_even(14)
assert 2450 == sum_smaller_even(100)
### BEGIN HIDDEN TEST
assert 2756 == sum_smaller_even(105)
assert 13340 == sum_smaller_even(231)
### END HIDDEN TEST
```

Compute the mathematical sum $$\sum_{k=1}^M \frac{1}{k},$$ as a function of $M$, up to $M=10000$. Print the results to a file and plot them (check `matplotlib/seaborn`). 

```python

```

```python

```

Write a function to compute, numerically, the central derivative of a given function, $$ f'(x) \simeq \frac{f(x+h) - f(x-h)}{2h}.$$ Use, as a model, the function $sin(x)$ (and later try with $e^x$ and $x^2$). Print a table showing the difference between the exact derivate and the numerical approximation, as a function of $h$, and save it to a file. Change $h$ logarithmically (powers of 10), starting at $h = 0.1$ until $h=10^{-6}$. Evaluate the derivative at $x = 0.1$.  

```python

```

```python

```